Birds migrate longitudinally in response to the resultant Asian monsoons of the Qinghai-Tibet Plateau uplift

Author response:

The following is the authors’ response to the previous reviews

Reviewer #1 (Public review):

The authors have done a good job of responding to the reviewer's comments, and the paper is now much improved.

Again, we thank the reviewer for positive comments during review.

Reviewer #2 (Public review):

I would like to thank the authors for the revision and the input they invested in this study.

We are grateful for your thoughtful feedback and enthusiasms, which helps us improve our manuscript.

With the revised text of the study, my earlier criticism holds, and your arguments about the counterfactual approach are irrelevant to that. The recent rise of the counterfactual approach might likely mirror the fact that there are too many scientists behind their computers, and few go into the field to collect in situ data. Studies like the one presented here are a good intellectual exercise but the real impact is questionable.

We understand your concern about the relevance of the counterfactual approach used in our study. Our intent in using a counterfactual scenario (reconstructing migration patterns assuming pre-uplift conditions on the QTP) was to isolate the potential influence of the plateau’s geological history on current migration routes. Similar approach was widely used to estimate how biogeographic barriers facilitated the divergent vertebrate communities across the world (e.g., Williams et al. 2024). We agree that such an approach must be used carefully. In the revision, we have explicitly clarified why this counterfactual comparison is useful – namely it provides a theoretical baseline to test how much the QTP’s uplift (and the associated monsoon system) might have redirected migration paths (Gilbert and Lambert 2010, Sanmartín 2012, Bull et al. 2021). We acknowledge that the counterfactual results are theoretical and have explicitly emphasised the assumptions involved (i.e., species–environment relationships hold between pre- and post- lift environments) in the main text (Lines 91- 98). Nonetheless, we defend the approach as a valuable study design: it helps generate testable hypotheses about migration (for instance, that the plateau’s monsoon-driven climate, rather than just its elevation, introduces an east–west shift en route).

References:

Bull, J. W., N. Strange, R. J. Smith, and A. Gordon. 2021. Reconciling multiple counterfactuals when evaluating biodiversity conservation impact in social-ecological systems. Conservation Biology 35:510-521.

Gilbert, D., and D. Lambert. 2010. Counterfactual geographies: worlds that might have been. Journal of Historical Geography 36:245-252.

Sanmartín, I. 2012. Historical Biogeography: Evolution in Time and Space. Evolution: Education and Outreach 5:555-568.

Williams, P. J., E. F. Zipkin, and J. F. Brodie. 2024. Deep biogeographic barriers explain divergent global vertebrate communities. Nature Communications 15:2457.

All your main conclusions are inferred from published studies on 7! bird species. In addition, spatial sampling in those seven species was not ideal in relation to your target questions. Thus, no matter how fancy your findings look, the basic fact remains that your input data were for 7 bird species only! Your conclusion, “our study provides a novel understanding of how QTP shapes migration patterns of birds” is simply overstretching.

We appreciate the reviewer’s comment here. We would like to clarify that our conclusions regarding longitudinal shifts in migratory distributions are based on distribution models derived from eBird data of 50 species, not merely on migration tracks from seven species. These species-level spatiotemporal models allow us to infer large-scale biogeographic patterns across the Qinghai-Tibet Plateau (QTP).

The original seven tracking species were used specifically for analysing the relationship between migration directions (azimuths) and environmental variables, offering independent support for the patterns revealed in the eBird-based distribution models. Recognising the reviewer’s concern on sample size and coverage, we have now expanded this part by incorporating migration tracks from 12 additional species, derived through georeferenced digitisation of published migratory maps. Importantly, this expansion did not change our conclusions, i.e., the monsoons instead of the high elevations act as a prominent role in shaping the current migration direction of birds in the QTP. While the overall conclusion remains unchanged, the expanded dataset led to slight changes in difference between spring and autumn migration. We have updated the Figure 2 and the corresponding results and conclusions throughout the manuscript. We have also clarified in the Discussion that regions of the QTP with relatively less data might lead to underestimation of some migration routes to make sure readers are aware of these data limitations (Lines 211-218).

The way you respond to my criticism on L 81-93 is something different than what you admit in the rebuttal letter. The text of the ms is silent about the drawbacks and instead highlights your perspective. I understand you; you are trying to sell the story in a nice wrapper. In the rebuttal you state: “we assume species' responses to environments are conservative and their evolution should not discount our findings.” But I do not see that clearly stated in the main text.

Thanks, as suggested we have clearly stated the assumptions of niche conservatism in the Introduction (Lines 91-98).

In your rebuttal, you respond to my criticism of "No matter how good the data eBird provides is, you do not know population-specific connections between wintering and breeding sites" when you responded: ... "we can track the movement of species every week, and capture the breeding and wintering areas for specific populations" I am having a feeling that you either play with words with me or do not understand that from eBird data nobody will be ever able to estimate population-specific teleconnections between breeding and wintering areas. It is simply impossible as you do not track individuals. eBird gives you a global picture per species but not for particular populations. You cannot resolve this critical drawback of your study.

We agree that inferring population-specific migratory connections (teleconnections) from eBird data is challenging and inherently limited. eBird provides occurrence records for species, but it generally cannot distinguish which breeding population an individual bird came from or exactly where it goes for winter. Our objective is not to determine one-to-one migratory links between specific populations, but to identify general broad-scale directional shifts when birds cross the QTP during their migration. We regret any confusion caused by our earlier wording. To make this clearer, we have now emphasised that our interests focus on the migratory direction and their environmental correlates, rather than population assignments. We have also rephrased the relevant text to explicitly clarify that our study operates at the species level and at large spatial scales (Lines 253–257). We exemplify how distribution of eBird observations and GPS tracking data of four species can be different from each other whilst showing similar migration patterns (Figure S10). We have also explicitly stated in the Discussion that confirming population connectivity would require targeted tracking or genetic studies, and that our eBird-based analysis could only suggest plausible routes and region-to-region linkages (Lines 200-202).

I am sorry that you invested so much energy into this study, but I see it as a very limited contribution to understanding the role of a major barrier in shaping migration.

We thank the reviewer’s honest assessment and understand the concern regarding the scope of our contribution. Our intention was not to provide an exhaustive account of all aspects of the QTP as a migratory barrier, but to address a specific and underexplored question: how the uplift of the plateau and the resulting monsoon system may have influenced the orientation of avian migration routes. By integrating both satellite tracking and community-contributed data, we have explored how the uplift of the QTP could shape avian migration across the area. We believe our findings provide important insights of how birds balance their responses to large-scale climate change and geological barrier, which yields the most comprehensive picture to date of how the QTP uplift have shaped migratory patterns of birds. We have also discussed the study’s limitations – including the small number of tracking species (Lines 205218), the use of occurrence data as a proxy for breeding and wintering regions (Lines 200-202), the uneven sampling coverage in the QTP (Lines 202-205) and the assumptions behind the counterfactual scenario (Lines 91-98). This ensures that readers understand the context and constraints of our findings.

My modest suggestion for you is: go into the field. Ideally use bird radars along the plateau to document whether the birds shift the directions when facing the barrier.

We thank the reviewer for this suggestion. We agree that radar holds promise for understanding certain aspects of bird migration, particularly for detecting flight intensity, altitudes, and timing. However, the radar systems are currently challenging to resolve migration at the level of species, populations, or individuals, which are central to questions of migratory connectivity and route selection. Most radar signals cannot distinguish between species in mixed flocks, nor can they link breeding and wintering sites for tracked individuals. In addition, the spatial coverage of radar installations remains limited, especially across remote and high-elevation regions like the Qinghai-Tibet Plateau, where infrastructure and continuous power supply are still logistically prohibitive.

The eBird dataset used in our study is itself a form of field-based observation, contributed by tens of thousands of birdwatchers across continents, including the QTP region (Figure S11). While eBird cannot provide individual-level tracking, it captures spatiotemporal patterns of occurrence at broad scales, making it a valuable complement to satellite tracking data. We would also emphasis that our team has extensive field experience in the Qinghai-Tibet Plateau (about twenty years), including multi-year expeditions to deploy satellite tags and observe migration at stopover sites.

We agree that more direct tracking (e.g. GPS tagging) would be an ideal way to validate migration pathways and population connectivity. Using the satellite-tracking data, we have showed that most tracking species shifted their migration direction when facing the QTP (Figure S6). In this revision, as stated we managed to add a number of 12 more species with satellite tracking routes. We have also noted that future studies should build on our findings by using dedicated tracking of more individual birds and monitoring of migration over the QTP. We have cited recent advances in these techniques and suggested that incorporating more tracking data could further test the hypotheses generated by our work (Lines 205-218).

Reviewer #2 (Recommendations for the authors):

L55 "an important animal movement behaviour is.." Is there any unimportant animal movement? I mean this sentence is floppy, empty.

We used this sentence to introduce migration. We have removed “important” to reduce ambiguous phrasing.

L 152-154 This sentence is full of nonsense or you misinterpretation. First of all, the issue of inflexible initiation of migration was related to long-distance migrants only! The way you present it mixes apples and oranges (long- and short-distance migrants). It is not "owing to insufficient responses" but due to inherited patterns of when to take off, photoperiod and local conditions.

We stated that this claim is invoked for long-distance migrants before this sentence and have rewritten the sentence to highlight that this interpretation is for long-distance migrants.

L 158 what is a migration circle? I do not know such a term.

We have amended it as “annual migration cycle”, which is a more common way to describe the yearly round-trip journey between breeding and wintering grounds of birds.

L 193 The way you present and mix capital and income breeding theory with your simulation study is quite tricky and super speculative.

We thank the reviewer for raising this important concern. We have presented this idea as an inference rather than a conclusion: “This pattern could be consistent with a ‘capital breeding’ strategy — where birds rely on endogenous reserved energy gained prior to reproduction — rather than an ‘income’ strategy where birds ingest nutrients mainly collected during the period of reproductive activity. This collaborates with studies on breeding strategies of migratory birds in Asian flyways. However, we note that this interpretation would require further study.” By adding this caution, we made it clear that we are not asserting this link as proven fact, only suggesting it as one possible explanation. We have also doublechecked that the rest of the discussion around this point is framed appropriately. Moreover, to help illustrate why we raised this ecological interpretation, we would also draw attention to examples of satellite tracking points from several species (e.g., Beijing Swift, Demoiselle Crane) in the following, which show obvious shifts in migratory direction near the QTP region. These turning points suggest potential behavioral responses to environmental constraints, such as climatic corridors or energy availability, which could help motivate our discussion of possible capital breeding strategies in these species.

Changes from the earlier version to the revised version

Abstract

Abstract

The uplift of the Qinghai-Tibet Plateau is one of the greatest geological events on Earth, pivotally shaping biogeographic patterns across continents, especially for migratory species that need to overcome topographical barriers to fulfil their annual circle. However, how the uplift influences animal migration strategies remains largely unclear. We compare the current flyways of 50 avian species migrating across the plateau with those reconstructed before the uplift as a counterfactual. We find that the major effect of the plateau uplift is changing avian migratory directions from the latitudinal to the longitudinal. The monsoon system generated by the uplift rather than the high elevation per se shapes those changes. These findings unveil how an important global geological event has influenced biogeographic patterns of migratory birds, yielding testable hypotheses for how observed avian distributions emerge.

Introduction

The Qinghai-Tibet Plateau (QTP) is the most extensively elevated surface on Earth, with an average elevation of ∼5 km over an area of 2.5 million km2 [1]. The uplift of the plateau exerts profound influences on the environment and determines the biogeographic boundaries both within and across continents [2]. The unique geological developments of the QTP, especially its high elevation, are believed to have influenced various taxonomic groups continuously [3–5]. However, the plateau’s uplift also brings up Asian monsoons, one of the most vigorous phenomena in the global climate system [6,7]. The Asian monsoons dominate large areas extending from the Indian sub-continent eastwards to Southeast and East Asia [8]. Their evolution and variability have caused significant variations in the redistribution of water and heat via a series of natural processes, such as drought, flood, and heat waves [6]. Given the large impacts of Asian monsoons on climate and environments, they can reconfigure the spatial patterns of biodiversity and ecosystem processes. This entails movement patterns that shape the effects of the environment on organisms [9–11]. However, owing to the difficulties in studying the complex effects caused by monsoons, most studies that explored the influence of the QTP just conceived the plateau as an orographic barrier [12–14]. The role of monsoons in shaping species movement patterns remains poorly understood.

Earlier version: Animal movement underpins species’ spatial distributions and ecosystem processes. An important animal movement behaviour is migrating between breeding and wintering grounds [15–17]. Those migratory journeys have intrigued a body of different approaches and indicators to describe and model migration, including migratory direction, speed, timing, distance, and staging periods [18,19]. Amongst them, the migratory direction is one of the most prominent indicators for migration patterns, evidenced by a majority of animals migrating latitudinally between wintering and breeding areas [20]. This can be explained by not only the fact that wintering sites are usually located in the warmer south (e.g. tropic) and breeding sites located in the cooler north (e.g. arctic), but also the earth’s magnetic fields that are arguably believed to affect the latitudinal migration of animals [21–24]. However, the migratory direction can be changed from latitudinal to longitudinal when the animal faces environmental changes [20,25–28].

Revised version: Animal movement underpins species’ spatial distributions and ecosystem processes. An animal movement behaviour is migrating between breeding and wintering grounds [15–17]. Those migratory journeys have intrigued a body of different approaches and indicators to describe and model migration, including migratory direction, speed, timing, distance, and staging periods [18,19]. Amongst them, the migratory direction is one of the most prominent indicators for migration patterns, evidenced by a majority of animals migrating latitudinally between wintering and breeding areas [20]. This can be explained by not only the fact that wintering sites are usually located in the warmer south (e.g. tropic) and breeding sites located in the cooler north (e.g. arctic), but also the earth’s magnetic fields that are arguably believed to affect the latitudinal migration of animals [21–24]. However, the migratory direction can be changed from latitudinal to longitudinal when the animal faces environmental changes [20,25–28].

Environmental fluctuations in the QTP are relatively small over the longue durée after the final- stage uplift [29], but few studies have evaluated how environmental heterogeneity across the QTP might influence the migratory behaviour of birds (but see migratory pattern descriptions, e.g., Zhao et al. [30], Pu and Guo [31]). Yet it remains unclear whether and how these shifts systematically alter species’ migration patterns rather than a simple assumption that the QTP birds exploit resources according to their availability. Therefore, testing whether migration patterns vary consistently for birds that migrate across the QTP is key to our understanding of the processes that determine movement patterns and provides insights into how they may affect community organisation and functioning under the context of global environmental change.

Earlier version: In this work, we leverage the community-contributed and satellite-tracking data to explore the impacts of the QTP uplift in terms of both the development of its high elevation and Asian monsoons on the migratory strategies for the birds that migrated across the plateau. We do this by reconstructing the environments before the uplift and contrasting migratory directions of 50 bird species (See Table S1 for a full list of species) between breeding and wintering areas in environments before the uplift with those at present. Thus, the simulated environments before the uplift of the plateau serve as a counterfactual state. Counterfactual is an important concept to support causation claims by comparing what happened to what would have happened in a hypothetical situation: “If event X had not occurred, event Y would not have occurred” [32]. Recent years have seen an increasing application of the counterfactual approach to detect biodiversity change, i.e., comparing diversity between the counterfactual state and real estimates to attribute the factors causing such changes e.g., Gonzalez et al. [33]. Whilst we do not aim to provide causal inferences for avian distributional change, using the counterfactual approach, we are able to estimate the influence of the plateau uplift by detecting the changes of avian distributions, i.e., by comparing where the birds would have distributed without the plateau to where they currently distributed. We regard the counterfactual environments as an ideal tool for eliminating, to the extent possible, vagueness, as opposed to simply description of current distributions of birds.

Revised version: In this work, we leverage the community-contributed and satellite-tracking data to explore the impacts of the QTP uplift in terms of both the development of its high elevation and Asian monsoons on the migratory strategies for the birds that migrated across the plateau. We do this by reconstructing the environments before the uplift and contrasting migratory directions of 50 bird species (See Table S1 for a full list of species) between breeding and wintering areas in environments before the uplift with those at present. Thus, the simulated environments before the uplift of the plateau serve as a counterfactual state. Counterfactual is an important concept to support causation claims by comparing what happened to what would have happened in a hypothetical situation: “If event X had not occurred, event Y would not have occurred” [32]. Recent years have seen an increasing application of the counterfactual approach to detect biodiversity change, i.e., comparing diversity between the counterfactual state and real estimates to attribute the factors causing such changes e.g., Gonzalez et al. [33]. For example, counterfactual is typically needed to evaluate the effectiveness of conservation interventions [34], and it helps identify the role of biogeographic barrier in shaping the diversity of global vertebrates [35]. Whilst we do not aim to provide causal inferences for avian distributional change, using the counterfactual approach, we are able to estimate the influence of the plateau uplift by detecting the changes of avian distributions, i.e., by comparing where the birds would have distributed without the plateau to where they currently distributed. A default assumption in the counterfactual analysis that should be carefully considered is that the species’ responses to environmental change (i.e., pre- and post- the uplift of the QTP in our analysis) are conservative, and only in this way could we test the influence of changed environments on species. Here, we regard the counterfactual environments as an ideal tool to generate testable hypotheses on the role of the QTP uplift, because it allows for isolating the potential influence of the plateau’s geological history on current migration routes to eliminate, to the extent possible, vagueness, as opposed to simply description of current distributions of birds.

We also calculate the migratory directions (azimuths) between adjacent stopover sites, breeding and wintering areas en route, and assess the relationship between migratory directions and environmental stress. Our findings yield the most comprehensive picture to date of how the QTP uplift shapes migratory patterns of birds, revealing insights into the challenges and opportunities for migratory birds in a changing world.

Earlier version: Results and discussion

Revised version: Results and discussions

Earlier version: We have two major findings regarding distribution patterns and migratory directions of QTP birds. First, we developed a dynamic species distribution modelling [18] to track the weekly distribution of target species, capturing the interconnections of stopover, wintering and breeding areas (See methods for details). By contrasting their distributions before and after the uplift, we find the distribution of migratory birds extended in longitude and narrowed in latitude with the uplift of QTP (Figure 1-E, G, and I). Birds are more likely to migrate along a longitudinal gradient in present environments as a result of the QTP uplift (See Table S1 for AUC values for model performance of each species). Specifically, before the uplift, migratory birds have a higher probability of breeding across a vast area at low and middle latitudes on the Eurasia continent, including West Asia, Siberia, QTP regions, and even Africa, whilst their most likely breeding areas move northeastward to the extreme north of Russia after the uplift. Different from the breeding area, the wintering area of migratory birds has a larger change in distribution probability. Birds that migrate across the QTP in the modern scenario have a higher probability of wintering in Southwest Asia and North Africa, whereas they have a higher probability of moving southeast to winter in Southern China and more areas of Africa before the uplift (Figure 1-D F and H)..

Revised version: We have two major findings regarding distribution patterns and migratory directions of QTP birds. First, we developed a dynamic species distribution modelling [18] to track the weekly distribution of target species, capturing the interconnections of stopover, wintering and breeding areas (See methods for details). By contrasting their distributions before and after the uplift, we find the distribution of migratory birds extended in longitude and narrowed in latitude with the uplift of QTP (Figure 1-E, G, and I). Birds are more likely to migrate along a longitudinal gradient in present environments as a result of the QTP uplift (See Table S1 for AUC values for model performance of each species). Specifically, before the uplift, migratory birds have a higher probability of breeding across a vast area at low and middle latitudes on the Eurasia continent, including West Asia, Siberia, QTP regions, and even Africa, whilst their most likely breeding areas move northeastwards to the extreme north of Russia after the uplift. Different from the breeding area, the wintering area of migratory birds has a larger change in distribution probability. Birds that migrate across the QTP in the modern scenario have a higher probability of wintering in Southwest Asia and North Africa, whereas they have a higher probability of moving southeast to winter in Southern China and more areas of Africa before the uplift (Figure 1-D F and H)..

Influence of the Qinghai-Tibet uplift on avian migration strategies.

(A) – (C) Schematic example of the role of Qinghai-Tibet Plateau (QTP) uplift in distribution patterns of migratory birds. (A) Birds migrate with a large longitudinal range in modern environments. Before the QTP uplift, birds may maintain similar migratory patterns with large longitudinal changes (B) or migrate with few longitudinal changes between wintering and breeding areas (C). The occurrence probability of 50 migratory bird species under modern environments in breeding areas (D) and wintering areas (E). The occurrence probability of birds in breeding areas (F) and wintering areas (G) before the QTP uplift. Migratory directions are identified at present (H) and before the uplift (I). The direction and length of the arrow represent migratory direction (measured by the azimuth angle) and distance from centres of breeding to wintering areas for each species. The circular barplot of the inset panel denotes the summary of migratory directions from breeding to wintering areas for each bird species, where the height and colour of the bars represent the number of species.

Earlier version: Second, our results show that wind cost, temperature, and precipitation are three major factors that influence the overall migratory directions (both autumn and spring) of birds, despite the differences in autumn and spring migration across different geographic areas (Figure 2). During autumn migration, wind cost is the most important factor for birds’ migration direction (Figure 2). A higher wind cost is associated with migration, which suggests a higher opportunity for birds to use the wind to facilitate their longitudinal migration (Figures S1 and S9). They also choose to follow a flyway of relatively higher annual precipitation and temperature as they migrate from breeding areas to wintering areas during autumn migration (Figure S2, S3 and S9). Apart from those three factors, no evidence is found for strong impacts of elevation and vegetation on the direction of migration (Figure 2 and Figure S9).

Revised version: Second, our results show that wind cost, temperature, and precipitation are three major factors that influence the overall migratory directions (both autumn and spring) of birds, despite the differences in autumn and spring migration across different geographic areas (Figure 2). Wind cost plays a more significant role during spring migration than that during autumn migration (Figure 2). A higher wind cost is associated with spring migration, which suggests a higher opportunity for birds to use the wind to facilitate their flight during spring migration (Figures S1 and S9). They choose to follow a flyway of relatively higher annual precipitation and temperature as they during both spring and autumn migration (Figure S2, S3 and S9). Apart from those three factors, no evidence is found for strong impacts of elevation and vegetation on the direction of migration (Figure 2 and Figure S9).

The influence of environmental factors on the direction of avian migration.

Migratory directions are calculated based on the azimuths between each adjacent stopover, breeding and wintering areas for each species. We employ multivariate linear regression models under the Bayesian framework to measure the correlation between environmental factors and avian migratory directions. Wind represents the wind cost calculated by wind connectivity. Vegetation is measured by the proportion of average vegetation cover in each pixel (∼1.9° in latitude by 2.5° in longitude). Temperature is the average annual temperature. Precipitation is the average yearly precipitation. All environmental layers are obtained using the Community Earth System Model. West QTP, central QTP, and East QTP denote areas in the areas west (longitude < 73°E), central (73°E ≤ longitude < 105°E), and east of (longitude ≥ 105°E) the Qinghai-Tibet Plateau, respectively.

Earlier version: Aside from the broad influences of QTP uplift, when migrating across different geographic areas, i.e., areas west of (longitude < 73°E, West QTP), areas in the central (73°E ≤ longitude < 105°E, middle QTP), and areas east of the QTP (longitude ≥ 105°E, East QTP), birds diversify their preferences in environmental conditions. Despite the fact that wind cost is the most important factor for the overall autumn migration, temperature is the most prominent factor in the areas east of the QTP and in the central QTP (Figure 2 and Figure S9). In addition, the average annual temperature in the central QTP is lower than that in the areas east of the QTP, but birds migrate across those two areas with increasing temperatures consistently (Figure S2, S6 and S9). Once they reach the regions west of the plateau (West QTP), low wind cost in the longitudinal direction and higher precipitation become priority choices for their migration (Figure 2 and Figure S9).

Revised version: Aside from the broad influences of QTP uplift, when migrating across different geographic areas, i.e., areas west of (longitude < 73°E, West QTP), areas in the central (73°E ≤ longitude < 105°E, middle QTP), and areas east of the QTP (longitude ≥ 105°E, East QTP), birds diversify their preferences in environmental conditions. Despite the fact that wind cost is the most important factor for the overall spring migration, temperature is the most prominent factor in the areas east of the QTP and in the central QTP (Figure 2 and Figure S9). Once they start to migration in the regions west of the plateau (West QTP), low wind cost in the longitudinal direction and higher precipitation are priority choices for their migration (Figure 2 and Figure S9). When they reach the central QTP, birds migrate across those areas with increasing temperatures (Figure S2, S6 and S9).

Earlier version: Compared with autumn migration, higher temperatures act as a major clue in the areas east and west of the plateau during spring migration, whereas the westerly outweighs temperature when birds migrate in the central plateau (Figure 1, Figure S1, S3 and S9). Besides temperature, precipitation also plays a role in all stages of spring migration. When birds migrate, they tend to follow a flyway of decreasing annual precipitation. Elevation has a slightly larger impact at the early stages of migration, i.e., areas east of the plateau during autumn migration and areas west of the plateau during spring migration, as birds migrate towards higher elevations during these stages.

Revised version: Compared with spring migration, higher temperatures act as a major clue in the areas east and west of the plateau during autumn migration, whereas the vegetation outweighs other factors when birds migrate in the central plateau (Figure 1, Figure S1, S3 and S9). Besides temperature, precipitation also plays a role in all stages of autumn migration. When birds migrate, they tend to follow a flyway of decreasing annual precipitation in central QTP and West QTP..

Earlier version: It is commonly claimed that the initiation of migration is inherently inflexible in migratory birds [34], owing to the weak or insufficient responses by migratory birds to adjusting migration behaviour (e.g., migration timing and route) [35]. This claim is particularly invoked for long-distance migrants, who may face greater temporal (e.g., migration timing) or physiological constraints given the varied phenologies en route [35]. Our results show that a major change in avian migratory patterns in response to environmental change can be adjusting migration direction from the latitudinal to the longitudinal at the scale of their whole migration circle. This highlights substantial changes in migratory bird distribution and their biogeographic patterns as a result of the uplift of the Qinghai-Tibet Plateau (Figure 1).

Revised version: It is commonly claimed that the initiation of migration is inherently inflexible in migratory birds [36], owing to the weak or insufficient responses by migratory birds to adjusting migration behaviour (e.g., migration timing and route) [37]. This claim is particularly invoked for long- distance migrants, who may face greater temporal (e.g., migration timing) or physiological constraints given the varied phenologies en route [37]. Our results show that a major change in avian migratory patterns in response to environmental change for long-distance migrants can be adjusting migration direction from the latitudinal to the longitudinal at the scale of their whole annual migration circle. This highlights substantial changes in migratory bird distribution and their biogeographic patterns as a result of the uplift of the Qinghai-Tibet Plateau (Figure 1).

Earlier version: One of the biggest climatic consequences of the uplift of the Qinghai-Tibet Plateau is the development of a unique monsoon system that has shaped environments across continents [36]. One typical feature of Asian monsoons is the seasonal climatic change, comprising a dry cold winter phase and a wet warm summer phase [36]. Asian monsoons also consist of several sub-systems, including the northeast monsoon and the East Asian winter monsoons that dominate the weather and climate in different parts of the plateau across different geographical periods. Our results showed that wind cost, temperature, and precipitation have more important impacts on avian migration than elevations in different geographic areas (Figure 2). This suggests that the monsoon system, rather than the high elevations of the plateau per se, is an important factor during avian migration on the plateau (Figure 2). Specifically, when birds begin their autumn migration in early September, the influence of the Siberian High on migration emerges as the East Asian winter monsoons start to reach the area east of the QTP, and their impacts at this stage are mainly reflected by varied temperatures and relatively less precipitation [37,38]. This can explain why higher temperatures and more precipitation play a more important role than wind cost in the area east of the plateau (Figure 2) since higher annual temperatures and more precipitations mean more food resources for migrants [39,40], whereas wind during this period is less strong than that in winter [38]. Whilst birds migrate westwards, less wind cost is becoming more important to determine their migration direction. This is, on the one hand, because the northeast monsoon begins to dominate the climate in the southwest of the QTP from around the end of September [41]. The northeast monsoon brings cold wind to sweep the Qinghai-Tibet Plateau down towards the vast spans of the Indian Ocean [41], which could facilitate the westward migration of birds. On the other hand, in the northwest of the Qinghai-Tibet Plateau, the extended Siberia High and associated atmospheric systems that deliver cold and dry air masses to the Mediterranean surface also provide positive wind conditions for migrants [42].

Revised version: One of the biggest climatic consequences of the uplift of the Qinghai-Tibet Plateau is the development of a unique monsoon system that has shaped environments across continents [38]. One typical feature of Asian monsoons is the seasonal climatic change, comprising a dry cold winter phase and a wet warm summer phase [38]. Asian monsoons also consist of several sub- systems, including the northeast monsoon and the East Asian winter monsoons that dominate the weather and climate in different parts of the plateau across different geographical periods. Our results showed that wind cost, temperature, and precipitation have more important impacts on avian migration than elevations in different geographic areas (Figure 2). This suggests that the monsoon system, rather than the high elevations of the plateau per se, is an important factor during avian migration on the plateau (Figure 2). Specifically, when birds begin their autumn migration in early September, the influence of the Siberian High on migration emerges as the East Asian winter monsoons start to reach the area east of the QTP, and their impacts at this stage are mainly reflected by varied temperatures and relatively less precipitation [39,40]. This can explain why higher temperatures and more precipitation play a more important role than wind cost in the east and central QTP (Figure 2) since higher annual temperatures and more precipitations mean more food resources for migrants [41,42], whereas wind during this period is less strong than that in winter [40]. Whilst birds migrate in the west QTP, less wind cost is becoming more important to determine their migration direction. This is, on the one hand, likely because the northeast monsoon begins to dominate the climate in the southwest of the QTP from around the end of September [43]. The northeast monsoon brings cold wind to sweep the Qinghai-Tibet Plateau down towards the vast spans of the Indian Ocean [43], which could facilitate the westward migration of birds. On the other hand, in the northwest of the Qinghai- Tibet Plateau, the extended Siberia High and associated atmospheric systems that deliver cold and dry air masses to the Mediterranean surface can also provide positive wind conditions for migrants [44].

Earlier version: When migrating toward breeding grounds in spring, birds adopt strategies different from their autumn migration, accompanied by different effects of environment on their migratory directions. Temperature becomes more important than wind cost in spring migration (Figure 2 and S9). Given high temperatures usually mean relatively rich food resources for birds [43,44], this suggests that birds that migrate across the Qinghai-Tibet Plateau may focus on energy accumulation during their spring migration rather than reducing flight costs in an effort to meet the energetic demands. Those birds also tend to follow a more ‘capital’ breeding strategy where birds rely on endogenous reserved energy gained prior to reproduction [45] rather than an ‘income’ breeding strategy where birds ingest nutrients mainly collected during the period of reproductive activity [46]. This collaborates with studies on breeding strategies of migratory birds in Asian flyways [47]. Another important reflection of the migration strategy is the role of elevation during different geographic areas of migration. As the elevation in the areas east and west of the plateau is much lower than that in the central, when migratory birds fly across the plateau, they need to follow a turned “u-shaped” elevation distribution and fly toward a higher elevation at the early and late stages of spring migration. This also implies that more energy is needed at the beginning of spring migration. Considering the need to balance energy accumulation and flight cost, areas with higher annual average temperatures and precipitation with high-level food resources are preferred during this migration stage to meet birds’ energy requests.

Revised version: When migrating towards breeding grounds in spring, birds adopt strategies different from their autumn migration, accompanied by different effects of environment on their migratory directions. Temperature is more important than wind cost in spring migration (Figure 2 and S9). Given high temperatures usually mean relatively rich food resources for birds [45,46], this suggests that birds that migrate across the Qinghai-Tibet Plateau may focus on energy accumulation during their spring migration rather than reducing flight costs in an effort to meet the energetic demands. This pattern could be explained by a ‘capital breeding’ strategy — where birds rely on endogenous reserved energy gained prior to reproduction [47] — rather than an ‘income’ strategy where birds ingest nutrients mainly collected during the period of reproductive activity [48]. This collaborates with studies on breeding strategies of migratory birds in Asian flyways [49]. However, we note that this interpretation would require further study.

Caveats and conclusions

Earlier version: Whilst we adopted both community-contributed and tracking data where potential biases existed, there are caveats to be aware of when interpreting our results. First, we used adaptive spatiotemporal modelling to address the imbalanced distribution of sampling in eBird data, but more sampling efforts and observations are still needed in areas of sparse records to better model and predict changes of species distributions. Second, tracking data can provide detailed information of the movement patterns of species but are limited to small numbers of species due to the considerable costs and time needed. We aimed to adopt the tracking data to examine the influence of focal factors on avian migration patterns, but only seven species, to the best of our ability, were acquired. Similar results were found in studies that used tracking data to estimate the distribution of breeding and wintering areas of birds in the plateau [e.g., 30,31,48-54]. The results based on seven species are rigour, but their implications could be restricted by the number of tracking species we obtained. We call for more tracking data to test and investigate the influence of QTP on multiple aspects of avian migratory patterns.

Revised version: Whilst we adopted both community-contributed and tracking data where potential biases existed, there are caveats to be aware of when interpreting our results. First, we adopted eBird data to infer plausible broad-scale region-to-region distributions of birds at the species level but confirming population connectivity would require targeted tracking or genetic studies. Second, we used adaptive spatiotemporal modelling to address the imbalanced distribution of sampling in eBird data, but more sampling efforts and observations are still needed in areas of sparse records to better model and predict changes of species distributions. Third, tracking data can provide detailed information of the movement patterns of species but are limited to small numbers of species due to the considerable costs and time needed. We aimed to adopt the tracking data to examine the influence of focal factors on avian migration patterns, but only 19 species, to the best of our ability, were acquired. Similar results were found in studies that used tracking data to estimate the distribution of breeding and wintering areas of birds in the plateau [e.g., 30,31,50- 56]. The results based on 19 species are rigour, but their implications could be restricted by the number of tracking species we obtained, e.g., limited number of tracking individuals might lead to underestimation of some migration routes. Recent years have seen an increasing number of tracking techniques which enable us to acquire accurate information on the movements of multiple individual animals in the wild, including satellite tracking, computer vision, radar and geolocator [57]. Future studies should build on our findings by using dedicated tracking of more individual birds and radar monitoring of animal migration over the QTP to test and investigate the influence of QTP on multiple aspects of avian migratory patterns.

Despite these caveats, our study provides a novel understanding of how QTP shapes migration patterns of birds. Albeit with the extensive influence of the plateau uplift on geology and geography, the resultant monsoon system, rather than its high elevation, is found to be a key factor shaping present avian migration patterns. Our study unveils shifts in avian migratory directions and their underlying mechanisms in the contexts of the QTP uplift, enhancing comprehension of the complex biogeographic effects on animal migration.

Methods and materials

Summary

Earlier version: We used two approaches to determine the migratory flyways of birds across the Qinghai-Tibet Plateau. First, we quantified the distributional change of each avian species by comparing the distribution range before and after the uplift of the plateau. For the present distribution, we used a dynamic spatiotemporal abundance model - Adaptive Spatiotemporal Model (AdaSTEM) that we have developed - to obtain the seasonal distribution of birds [18]. We then used a species distribution model (i.e., MaxEnt) to measure the correlation between present distribution and environments [55]. We calculated the distribution of migratory birds before the uplift of the plateau by projecting the correlation between their current distribution and environments onto environments before the uplift. Second, we obtained the specific migratory routes for each species by measuring the migratory directions (i.e., the azimuth angle between adjacent stopover sites and breeding and wintering areas) en route. Similarly, we used the relationships between present migratory directions and environments to predict the migratory directions pre-uplift of the plateau. Since our aim here was a prediction, we used random forest models, but we also used Bayesian multivariate regression modelling to measure the influence of environments on migratory directions of birds.

Revised version: We used two approaches to determine the migratory flyways of birds across the Qinghai-Tibet Plateau. First, we quantified the distributional change of each avian species by comparing the distribution range before and after the uplift of the plateau. For the present distribution, we used a dynamic spatiotemporal abundance model - Adaptive Spatiotemporal Model (AdaSTEM) that we have developed - to obtain the seasonal distribution of birds [18]. We then used a species distribution model (i.e., MaxEnt) to measure the correlation between present distribution and environments [58]. We calculated the distribution of migratory birds before the uplift of the plateau by projecting the correlation between their current distribution and environments onto environments before the uplift. Second, we obtained the specific migratory routes for each species by measuring the migratory directions (i.e., the azimuth angle between adjacent stopover sites and breeding and wintering areas) en route. Similarly, we used the relationships between present migratory directions and environments to predict the migratory directions pre-uplift of the plateau. Since our aim here was a prediction, we used random forest models, but we also used Bayesian multivariate regression modelling to measure the influence of environments on migratory directions of birds.

eBird checklist

Earlier version: We used a community-contributed database for the dynamic spatiotemporal abundance model to measure the seasonal distribution. Specifically, we first obtained the list of bird species that might migrate across the Qinghai-Tibet plateau based on Prins and Namgail [56]. We then requested and downloaded the eBird Basic Dataset in Feb 2022 [57] for 64 species. We then excluded species that were not listed as “full migrant” in BirdLife International (https://datazone.birdlife.org), which resulted in a total of 50 avian species analysed in our study and covered breeding populations in geographical Asia. We used data from the year 2019 to avoid the potential influence of the pandemic on bird observation [58,59] and bird behaviour [60].

Revised version: We used a community-contributed database for the dynamic spatiotemporal abundance model to measure the seasonal distribution. Specifically, we first obtained the list of bird species that might migrate across the Qinghai-Tibet plateau based on Prins and Namgail [59]. We then requested and downloaded the eBird Basic Dataset in Feb 2022 [60] for 64 species. We then excluded species that were not listed as “full migrant” in BirdLife International (https://datazone.birdlife.org), which resulted in a total of 50 avian species analysed in our study and covered breeding populations in geographical Asia. We used data from the year 2019 to avoid the potential influence of the pandemic on bird observation [61,62] and bird behaviour [63]. It is worth noting that eBird provides occurrence records for species, but it generally cannot distinguish which breeding population an individual bird came from or exactly where it goes for winter. In this study, we used eBird data to infer broad-scale movement patterns (e.g. general directions and stopover regions) at the species level rather than precise one-to-one population linkages.

Earlier version: The eBird data may be biased by the imbalanced sampling and variation of observers’ skills in identifying species. To address spatiotemporal imbalances in data distribution and the potential overrepresentation of birding hotspots, we conducted spatiotemporal subsampling following the method proposed [61,62]. We first assigned each checklist with a global hexagonal hierarchical geospatial indexing system [H3 system; 63,64], with a resolution of level 7 (∼5.16 km per cell). Then, to avoid biased sampling in rare species with unusual active temporal periods, we split the 24 hours of the day into 12 equal bins and assigned a checklist to each of the bins. We then randomly subsampled only one checklist for each year - day of the year - hour bin of day - cell combination. The subsampling resulted in 5,037,088 checklists for the year 2019.

Revised version: The eBird data may be biased by the imbalanced sampling and variation of observers’ skills in identifying species. To address spatiotemporal imbalances in data distribution and the potential overrepresentation of birding hotspots, we conducted spatiotemporal subsampling following the method proposed [64,65]. We first assigned each checklist with a global hexagonal hierarchical geospatial indexing system [H3 system; 66,67], with a resolution of level 7 (∼5.16 km per cell). Then, to avoid biased sampling in rare species with unusual active temporal periods, we split the 24 hours of the day into 12 equal bins and assigned a checklist to each of the bins. We then randomly subsampled only one checklist for each year - day of the year - hour bin of day - cell combination. The subsampling resulted in 5,037,088 checklists for the year 2019.

Earlier version: To account for difference of observers’ expertise in recognising species,, we calculated the historical cumulative species count for each bird observer throughout their historical eBird checklists prior to 2019 as a proxy to measure the expertise of bird observers [65]. We then filtered the checklists as suggested by recent studies [61,62]:

Revised version: To account for difference of observers’ expertise in recognising species, we calculated the historical cumulative species count for each bird observer throughout their historical eBird checklists prior to 2019 as a proxy to measure the expertise of bird observers [68]. We then filtered the checklists as suggested by recent studies [64,65]:

Only checklists labelled as complete were included.

Only checklists with Traveling or Stationary protocol types were included. For checklists with the protocol type Traveling, only those with a travelling distance of less than 3km were included.

The observation duration should be longer than 5 min and shorter than 300 min.

Observers with expertise lower than 2.5% percentile were removed since they are less representative and may induce large bias.

Predictor variables for spatiotemporal abundance modelling

For each remaining checklist, we extracted six types of environmental variables based on their geographical coordinates:

Sampling effort variables, which include protocol type, travelling distance, observation duration minutes, number of observers, and observers’ expertise (measured in historical species count).

Temporal variables, which include day of year and observation started time of day.

Earlier version: Topographic variables, which include the mean and standard deviation of elevation, slope, north, and east aggregated within the 3 km × 3 km buffered area for each checklist. The Topographic data was downloaded from EarthEnv [66] in a 1 km resolution.

Revised version: Topographic variables, which include the mean and standard deviation of elevation, slope, north, and east aggregated within the 3 km X 3 km buffered area for each checklist. The Topographic data was downloaded from EarthEnv [69] in a 1 km resolution.

Earlier version: Land cover data. We used the Copernicus Climate Change Service (C3S) Global Land Cover data with a 300 m resolution [67]. We calculated the landscape variables for each of the land cover types presented in the 3 km × 3 km buffered area for each checklist, including percentage cover, patch density, largest patch index, edge density, mean patch size, standard deviation of patch size for each land cover type, and entropy across heterogeneous land cover patches.

Revised version: Land cover data. We used the Copernicus Climate Change Service (C3S) Global Land Cover data with a 300 m resolution [70]. We calculated the landscape variables for each of the land cover types presented in the 3 km X 3 km buffered area for each checklist, including percentage cover, patch density, largest patch index, edge density, mean patch size, standard deviation of patch size for each land cover type, and entropy across heterogeneous land cover patches.

Earlier version: Bioclimate variables. We downloaded the ERA-5 hourly data at a 0.25° resolution [68]. Hourly data of a 2 m temperature and total precipitation layer were firstly aggregated to daily level by taking the average. The day-level data were calculated using 19 bioclimate variables, which were then assigned to each checklist according to the geographical coordinates.

Revised version: Bioclimate variables. We downloaded the ERA-5 hourly data at a 0.25° resolution [71]. Hourly data of a 2 m temperature and total precipitation layer were firstly aggregated to daily level by taking the average. The day-level data were calculated using 19 bioclimate variables, which were then assigned to each checklist according to the geographical coordinates.

Earlier version: Normalized Difference Vegetation Index (NDVI). NDVI data were extracted from Terra Moderate Resolution Imaging Spectroradiometer (MODIS) Vegetation Indices 16-Day (MOD13A2) Version 6.1 product with a resolution of 16 days and 1 km [69]. We further aggregated the data to hexagon level 5 based on the H3 indexing system (with edge length ∼9.85km). For each hexagon, we leveraged the pyGAM package [70] to apply a GAM model with 30 splines to interpolate the data to temporal ranges that were not provided by the original data. This resulted in a daily resolution dataset. We calculated six features based on NDVI and included them in subsequent modelling, i.e., the median, maximum, and minimum of NDVI, and the median, maximum, and minimum of the first derivative of NDVI against day of year (sometimes referred to “green wave”) for each hexagon throughout the year.

Revised version: Normalized Difference Vegetation Index (NDVI). NDVI data were extracted from Terra Moderate Resolution Imaging Spectroradiometer (MODIS) Vegetation Indices 16-Day (MOD13A2) Version 6.1 product with a resolution of 16 days and 1 km [72]. We further aggregated the data to hexagon level 5 based on the H3 indexing system (with edge length ∼9.85km). For each hexagon, we leveraged the pyGAM package [73] to apply a GAM model with 30 splines to interpolate the data to temporal ranges that were not provided by the original data. This resulted in a daily resolution dataset. We calculated six features based on NDVI and included them in subsequent modelling, i.e., the median, maximum, and minimum of NDVI, and the median, maximum, and minimum of the first derivative of NDVI against day of year (sometimes referred to “green wave”) for each hexagon throughout the year.

The feature engineering resulted in 106 predictor variables, including 6 sampling effort variables, 2 temporal variables, 8 topographic variables, 19 annual climatic variables, 65 land cover variables, and 6 vegetation index-related variables. All calculations are conducted in Python version 3.9.0.

Spatiotemporal abundance modelling

To adjust for sampling error and obtain the general migration pattern incorporating interconnections of stopover, wintering and breeding areas across species, we applied an Adaptive Spatio-Temporal Model (AdaSTEM) for each species to model weekly distributions of birds using stemflow package version 1.0.9.1, which we have recently reported [18].

Earlier version: AdaSTEM is a machine learning modelling framework that takes space, time, and sample size into consideration at different scales. It has been frequently used in modelling eBird data [62,71,72] and has been proven to be efficient and advanced in multi-scale spatiotemporal data modelling. To briefly summarise the methodology, in the training procedure, the model recursively splits the input training data into smaller spatiotemporal grids (stixels) using the QuadTree algorithm [73]. For each of the stixels, we trained a base model only using data contained by itself. Stixels were then aggregated and constituted an ensemble. In the prediction phase, stemflow queries stixels for the input data according to their spatial and temporal indexes, followed by the prediction of corresponding base models. Finally, we aggregated prediction results across ensembles to generate robust estimations (see Fink et al., 2013 [71] and stemflow documentation [18] for details).

Revised version: AdaSTEM is a machine learning modelling framework that takes space, time, and sample size into consideration at different scales. It has been frequently used in modelling eBird data [65,74,75] and has been proven to be efficient and advanced in multi-scale spatiotemporal data modelling. To briefly summarise the methodology, in the training procedure, the model recursively splits the input training data into smaller spatiotemporal grids (stixels) using the QuadTree algorithm [76]. For each of the stixels, we trained a base model only using data contained by itself. Stixels were then aggregated and constituted an ensemble. In the prediction phase, stemflow queries stixels for the input data according to their spatial and temporal indexes, followed by the prediction of corresponding base models. Finally, we aggregated prediction results across ensembles to generate robust estimations (see Fink et al., 2013 [74] and stemflow documentation [18] for details).

Earlier version: We used XGBoost [74] as our classifier and regressor base model for its capability and balance between performance and computational efficiency. We set 10 ensemble folds, a maximum grid length threshold of 25 degrees, a minimum grid length threshold of 5 degrees, a temporal sliding window size of 50 DOY and a step of 20 DOY, and required at least 50 checklists for each stixel in model training. Trained models were then used to predict on prediction dataset with 0.1° spatial resolution and weekly temporal resolution, where the variables were annotated with the same methodology as that of the training dataset. Only spatiotemporal points with more than seven ensembles covered are predicted. In downstream analyses, we removed data points with abundance lower than 0.1 quantiles to obtain reliable predictions for each week.

Revised version: We used XGBoost [77] as our classifier and regressor base model for its capability and balance between performance and computational efficiency. We set 10 ensemble folds, a maximum grid length threshold of 25 degrees, a minimum grid length threshold of 5 degrees, a temporal sliding window size of 50 DOY and a step of 20 DOY, and required at least 50 checklists for each stixel in model training. Trained models were then used to predict on prediction dataset with 0.1° spatial resolution and weekly temporal resolution, where the variables were annotated with the same methodology as that of the training dataset. Only spatiotemporal points with more than seven ensembles covered are predicted. In downstream analyses, we removed data points with abundance lower than 0.1 quantiles to obtain reliable predictions for each week.

Environmental variables for species distribution modelling

Given the challenges in simulating environmental and climatic conditions before the uplift of the Qinghai-Tibet Plateau, we modelled the environments before and after the uplift with five variables, i.e., monthly wind (speed and direction), annual temperature, annual precipitation, elevation and annual vegetation.

Earlier version: In detail, following Zhang et al. [7], we used version 1.0.4 of the Community Earth System Model (CESM) coupled model with a dynamic atmosphere (CAM4), land (CLM4), ocean (POP2), and sea-ice (CICE4) components to simulate pre-uplift environments. CESM and its previous versions have been widely used in climate modelling, e.g., Meehl et al. [75] and are claimed to be capable of broadly reproducing the features of present-day climate [76]. For CAM4, there is a horizontal resolution of ∼1.9° in latitude by 2.5° in longitude and 26 layers in the vertical direction. POP2 adopts a finer grid and has a nominal 1° horizontal resolution (320 × 384 grid points, latitude by longitude) and 60 layers in the vertical direction. The land and sea-ice components share the same horizontal grids as the atmosphere and ocean components, respectively. In CLM4, multiple land surface types and plant functional types (PFTs) are contained within one grid, and CLM4 can be run in a dynamic vegetation mode to simulate natural vegetation, including trees, grass, and shrub plant functional types, e.g., Yu et al. [77] Qiu and Liu [78].

Revised version: In detail, following Zhang et al. [7], we used version 1.0.4 of the Community Earth System Model (CESM) coupled model with a dynamic atmosphere (CAM4), land (CLM4), ocean (POP2), and sea-ice (CICE4) components to simulate pre-uplift environments. CESM and its previous versions have been widely used in climate modelling, e.g., Meehl et al. [78] and are claimed to be capable of broadly reproducing the features of present-day climate [79]. For CAM4, there is a horizontal resolution of ∼1.9° in latitude by 2.5° in longitude and 26 layers in the vertical direction. POP2 adopts a finer grid and has a nominal 1° horizontal resolution (320 × 384 grid points, latitude by longitude) and 60 layers in the vertical direction. The land and sea- ice components share the same horizontal grids as the atmosphere and ocean components, respectively. In CLM4, multiple land surface types and plant functional types (PFTs) are contained within one grid, and CLM4 can be run in a dynamic vegetation mode to simulate natural vegetation, including trees, grass, and shrub plant functional types, e.g., Yu et al. [80] Qiu and Liu [81].

We initiated the modelling with two different scenarios, i.e., the actual elevation and a maximum elevation of 300m. We then used the same default preindustrial simulation for the two scenarios with a modern ice sheet, an atmospheric CO2 concentration of 280 ppmv, modern orbital parameters (the year 1950), modern solar constant (1,365 W/m2), other atmospheric greenhouse gas concentrations set to preindustrial values (CH4 and N2O set to 760 and 270 ppbv, respectively), and preindustrial aerosol conditions. We ran for 750 model years to ensure the combined atmospheric, ocean, and vegetation effects in response to the uplift of the plateau can be investigated.

Species distribution modelling

Earlier version: We used Maximum entropy (MaxEnt) models to compare the avian distributional change between pre- and post-uplift environments under the assumption that species tend to keep their ancestral ecological traits over time (i.e., niche conservatism). This indicates a high probability for species to distribute in similar environments wherever suitable. Particularly, considering birds are more likely to be influenced by food resources and vegetation distributions [79–81], and the available food and vegetation before the uplift can provide suitable habitats for birds [82], we believe the findings can provide valuable insights into the influence of the plateau rise on avian migratory patterns. Having said that, we acknowledge other factors, e.g., carbon dioxide concentrations [83], can influence the simulations of environments and our prediction of avian distribution. MaxEnt compares the environmental features at presence points to those of pseudo absences to discriminate the suitable area [84]. MaxEnt builds models using a generative approach and thus has an inherent advantage over a discriminative approach, especially when the amount of training data is small [84]. Due to its good performance compared to other species distribution modelling techniques, MaxEnt is widely used in the study of biogeography and conservation biology.

Revised version: We used Maximum entropy (MaxEnt) models to compare the avian distributional change between pre- and post-uplift environments under the assumption that species tend to keep their ancestral ecological traits over time (i.e., niche conservatism). This indicates a high probability for species to distribute in similar environments wherever suitable. Particularly, considering birds are more likely to be influenced by food resources and vegetation distributions [82–84], and the available food and vegetation before the uplift can provide suitable habitats for birds [85], we believe the findings can provide valuable insights into the influence of the plateau rise on avian migratory patterns. Having said that, we acknowledge other factors, e.g., carbon dioxide concentrations [86], can influence the simulations of environments and our prediction of avian distribution. MaxEnt compares the environmental features at presence points to those of pseudo absences to discriminate the suitable area [87]. MaxEnt builds models using a generative approach and thus has an inherent advantage over a discriminative approach, especially when the amount of training data is small [87]. Due to its good performance compared to other species distribution modelling techniques, MaxEnt is widely used in the study of biogeography and conservation biology.

Earlier version: We ran the MaxEnt model using default settings but with 1,000 iterations. For each model, we ran 20 bootstrap replications, and each time 75% of locations were selected at random as training samples, while the remaining 25% were used as validation samples. We applied Area Under the Curve (AUC) of the Receiver Operator Characteristic (ROC)assess the performances of the models (Table S1). AUC is a threshold-independent measurement for discrimination ability between presence and random points [84]. When the AUC value is higher than 0.75, the model was considered to be good [85,86].

Revised version: We ran the MaxEnt model using default settings but with 1,000 iterations. For each model, we ran 20 bootstrap replications, and each time 75% of locations were selected at random as training samples, while the remaining 25% were used as validation samples. We applied Area Under the Curve (AUC) of the Receiver Operator Characteristic (ROC)assess the performances of the models (Table S1). AUC is a threshold-independent measurement for discrimination ability between presence and random points [87]. When the AUC value is higher than 0.75, the model was considered to be good [88,89].

Migratory direction

To obtain the species list of birds that migrate across the Qinghai-Tibet Plateau with available tracking data, we checked Movebank (movebank.org) together with literature reporting avian migratory routes across the plateau. For those who did not upload their data to Movebank, we digitalised the routes. Specifically, we built a new geographic layer with the same coordinate systems of each reported route and matched the layer with the images of routes. We then delineate migratory routes on the new geographic layer where the geographic information of the routes was achieved. This resulted in seven representative species that migrated across the plateau.

Earlier version: We used the same environmental variables, except for wind, for the species distribution model to investigate the potential influence of environments on migratory directions. We calculated wind connectivity to account for the influence of wind, considering wind connectivity has been identified as a key factor driving avian flying patterns [87]. Since we aimed to investigate the migration patterns at large spatiotemporal scales, we measured the wind connectivity at a monthly resolution to enable analysis of seasonal differences. We adjusted the R package rWind for the computation. In detail, we replaced the default wind data from the Global Forecasting System atmospheric model with our monthly wind data from CESM as input. For both wind costs before and after the uplift of the plateau, we then calculate the movement cost from any starting cell to one of its eight neighbouring cells (Moore neighbourhood). This includes three parameters, i.e., wind speed at the starting cell, wind direction at the starting cell (azimuth), and the position of the target cell. A movement connectivity map was then determined after performing the default algorithms [88]. We reversed the values of cells on the connectivity map, as we aimed to investigate the influence of wind cost, whereas the map showed the importance of the cell to maintain connectivity.

Revised version: We used the same environmental variables, except for wind, for the species distribution model to investigate the potential influence of environments on migratory directions. We calculated wind connectivity to account for the influence of wind, considering wind connectivity has been identified as a key factor driving avian flying patterns [90]. Since we aimed to investigate the migration patterns at large spatiotemporal scales, we measured the wind connectivity at a monthly resolution to enable analysis of seasonal differences. We adjusted the R package rWind for the computation. In detail, we replaced the default wind data from the Global Forecasting System atmospheric model with our monthly wind data from CESM as input. For both wind costs before and after the uplift of the plateau, we then calculate the movement cost from any starting cell to one of its eight neighbouring cells (Moore neighbourhood). This includes three parameters, i.e., wind speed at the starting cell, wind direction at the starting cell (azimuth), and the position of the target cell. A movement connectivity map was then determined after performing the default algorithms [91]. We reversed the values of cells on the connectivity map, as we aimed to investigate the influence of wind cost, whereas the map showed the importance of the cell to maintain connectivity.

We used a random forest model and a multivariate linear regression model under the Bayesian framework to analyse the influence of environments on avian migratory directions. We first used the random forest model to measure the correlation between migratory directions and modern environments and predict the migratory direction before the uplift of the plateau. We then compared the influence between modern environments and environments before the uplift using a multivariate linear regression model under the Bayesian framework. We adopted two strategies for those two modelling approaches. First, we applied regression to different combinations of season-stage separately (seasons: spring, autumn; stages: overall, east QTP, central QTP, west QTP), resulting in eight regression models. Second, we additionally included species as random variables by applying hierarchical modelling, which also resulted in eight regression models.

Earlier version: All variables were standardised for comparison. All Bayesian models were conducted with PyMC version 5.5 [89] in Python version 3.9.0 environment. We used a NUTS sampler with a numpyro backend (jax.sample_numpyro_nuts) to run four chains, each with 30,00 tuning and 3,000 posterior chain sampling. We assessed the model convergence using potential scale reduction factor (Rhat) and effective sample size (ESS), where all parameters in all models met the criteria of Rhat < 1.03 and ESS > 400.

Revised version: All variables were standardised for comparison. All Bayesian models were conducted with PyMC version 5.5 [92] in Python version 3.9.0 environment. We used a NUTS sampler with a numpyro backend (jax.sample_numpyro_nuts) to run four chains, each with 30,00 tuning and 3,000 posterior chain sampling. We assessed the model convergence using potential scale reduction factor (Rhat) and effective sample size (ESS), where all parameters in all models met the criteria of Rhat < 1.03 and ESS > 400.